# Week 2 Assignment

## Problem 1. Google Form.

Complete this [survey](https://docs.google.com/forms/d/1gH-6Fz4-oynqSJkgkUX3rhAnri4aCRdiMTcHCiP6S3g/viewform).

## Problem 2. Values and Types.

The purpose of this problem is to give you a chance to perform the following tasks:

- Run an IPython/Jupyter notebook server,
- Import an `.ipynb` notebook file onto the IPython/Jupyter server,
- Understand how a typical assignment works,
- Understand assignment statements and `int`, `float`, and `str`,
- Download a notebook from the IPython/Jupyter server and save it locally, and
- Upload completed assignments onto Moodle.

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select `Kernel` → `Restart`) and then run all cells (in the menubar, select `Cell` → `Run All`).

Make sure you fill in any place that says `YOUR CODE HERE`.

Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`.

First, write your name and NetID below:

```python
NAME = 'WRITE YOUR NAME HERE'
NETID = 'WRITE YOUR NETID HERE'
```

In the following code cell, assign a value of `490` to an integer variable named `an_integer`.

<!-- grade_id: int_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

Don't worry about it yet if you don't understand what the following text and the code cell after the text mean. It will make sense after you learn functions next week.

In the following code cell, you will see two assert statements, and almost every assignment will have such tests. These are what are called [unit tests](https://en.wikipedia.org/wiki/Unit_testing). When you submit completed assignments, instructors will determine your assignment grades by using automated grading tools to perform unit testing on your code, so if you pass all the tests without an error, you get a perfect score.

It should be obvious to you that these tests are valuable hints! Study them. Write each piece of code so that it doesn't produce any errors. Writing the tests first, and *then* writing the actual code is actually a perfectly valid approach (preferred by many) to programming. It's called [test-driven development](https://en.wikipedia.org/wiki/Test-driven_development).

So, remember to learn what `assert` means, and study [nose](http://nose.readthedocs.org/en/latest/testing.html) if you want to get a good grade.

<!-- grade_id: int_test, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(an_integer, 490)
assert_equal(type(an_integer), int)
```

In the following code cell, assign a value of `3.14159` to a floating point variable named `a_floating_point`.

<!-- grade_id: float_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: float_test, points: 10, grade: true, solution: false -->
```python
assert_equal(a_floating_point, 3.14159)
assert_equal(type(a_floating_point), float)
```

In the following code cell, assign a string `'informatics'` to a string variable named `a_string`.

<!-- grade_id: str_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: str_test, points: 10, grade: true, solution: false -->
```python
assert_equal(a_string, 'informatics')
assert_equal(type(a_string), str)
```