Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).

Make sure you fill in any place that says YOUR CODE HERE.
Do not write your answer in anywhere else other than where it says YOUR CODE HERE.

First, write your name and NetID below:

```python
NAME = 'WRITE YOUR NAME HERE'
NETID = 'WRITE YOUR NETID HERE'
```

# Problem 6.1. Regular Expressions: Flights Data.

In this problem, we will use Python's regular expressions (regex) to write functions that validate an IATA code, a longitude or a latitude, and a aircraft tail number.

```python
import re
from nose.tools import assert_equal
```

In [Problem 5.2](../../Week5/assignment/requests.ipynb), we introduced [IATA codes](https://en.wikipedia.org/wiki/International_Air_Transport_Association_airport_code) and [airports.csv](http://stat-computing.org/dataexpo/2009/supplemental-data.html).

The airports.csv file looks like this:

```shell
$ head /home/data_scientist/data/airports.csv
"iata","airport","city","state","country","lat","long"
"00M","Thigpen ","Bay Springs","MS","USA",31.95376472,-89.23450472
"00R","Livingston Municipal","Livingston","TX","USA",30.68586111,-95.01792778
"00V","Meadow Lake","Colorado Springs","CO","USA",38.94574889,-104.5698933
"01G","Perry-Warsaw","Perry","NY","USA",42.74134667,-78.05208056
"01J","Hilliard Airpark","Hilliard","FL","USA",30.6880125,-81.90594389
"01M","Tishomingo County","Belmont","MS","USA",34.49166667,-88.20111111
"02A","Gragg-Wade ","Clanton","AL","USA",32.85048667,-86.61145333
"02C","Capitol","Brookfield","WI","USA",43.08751,-88.17786917
"02G","Columbiana County","East Liverpool","OH","USA",40.67331278,-80.64140639
```

Consier the columns `iata`, `lat`, and `long`. Consider a scenario, where we have to go through every values in these columns, and see if they are valid because perhaps we suspect there might be some bad, corrupt, or missing values. How would we validate, for example, if some value in the `iata` column indeed has the correct format?

## 1. Validate an IATA code.

In the following code cell, write a function that returns `True` if the IATA code is valid and `False` if not.

For our purposes, an IATA airport code
  - has exactly three alphanumeric characters,
  - is valid when all alphabetical characters are capital letters, and
  - is invalid when one or more alphabetical characters are lowercase letters.
  
See the test cases for examples of valid and invalid IATA codes.

<!-- grade_id: validate_airport_answer, points: 0, grade: false, solution: true -->
```python
def validate_airport(iata):
    '''
    Takes a string.
    Returns True if the string is a valid IATA code.
    Returns False if the string is an invalid IATA code.
    
    Parameters
    ----------
    iata: string
    
    Returns
    -------
    bool
    '''
    # YOUR CODE HERE
```

The following code cell will produce no error if your function is correct.

<!-- grade_id: validate_airport_test, points: 10, grade: true, solution: false -->
```python
assert_equal(validate_airport('000'), True)
assert_equal(validate_airport('A00'), True)
assert_equal(validate_airport('0B0'), True)
assert_equal(validate_airport('00C'), True)
assert_equal(validate_airport('999'), True)
assert_equal(validate_airport('1AZ'), True)
assert_equal(validate_airport('B2Y'), True)
assert_equal(validate_airport('CX3'), True)

assert_equal(validate_airport(''), False)
assert_equal(validate_airport('1'), False)
assert_equal(validate_airport('23'), False)
assert_equal(validate_airport('4567'), False)
assert_equal(validate_airport('8y9'), False)
assert_equal(validate_airport('0aB'), False)
assert_equal(validate_airport('gHI'), False)
```

## 2. Validate a longitude or a latitude.

In the following code cell, write a function that returns `True` if a latitude/longitude is in the correct format, and returns `False` if the format is incorrect.

For our purposes, we will assume that a longitude or a latitude
  - always has decimal digits,
  - does not have plus sign when it is positive, and
  - has a minus sign when it is negative.
  
See the test cases for examples.

<!-- grade_id: validate_long_lat_answer, points: 0, grade: false, solution: true -->
```python
def validate_long_lat(coordinate):
    '''
    Takes a string.
    Returns True if the string is a valid longitude/latitude.
    Returns False if the string is not a valid longitude/latitude.
    
    Parameters
    ----------
    coordiante: string
    
    Returns
    -------
    bool
    '''
    # YOUR CODE HERE
```

<!-- grade_id: validate_long_lat_test, points: 10, grade: true, solution: false -->
```python
assert_equal(validate_long_lat('0.0'), True)
assert_equal(validate_long_lat('999.9'), True)

assert_equal(validate_long_lat('31.953'), True)
assert_equal(validate_long_lat('31.95376'), True)
assert_equal(validate_long_lat('31.95376472'), True)

assert_equal(validate_long_lat('-89.2'), True)
assert_equal(validate_long_lat('-89.23450472'), True)

assert_equal(validate_long_lat('+31.95376472'), False) # no plus sign for positive numbers

assert_equal(validate_long_lat('a0.0'), False) # no alphabet letters
assert_equal(validate_long_lat('1b.0'), False) # no alphabet letters
assert_equal(validate_long_lat('2.c0'), False) # no alphabet letters
assert_equal(validate_long_lat('999.9z'), False) # no alphabet letters

assert_equal(validate_long_lat('1'), False) # assume there will be decimal digits
assert_equal(validate_long_lat('1.'), False) # assume there will be decimal digits
```

## 3. Validate a tail number (aircraft registration number).

Accoridng to [Wikipedia](https://en.wikipedia.org/wiki/Aircraft_registration),

```
In the United States, the registration number is commonly referred to as an "N" number, because all aircraft registered there have a number starting with the letter N. An alphanumeric system is used because of the large numbers of aircraft registered in the United States. An N-number may only consist of one to five characters, must start with a digit other than zero, and cannot end in a run of more than two letters. In addition, N-numbers may not contain the letters I or O, due to their similarities with the numerals 1 and 0.
```

See the test cases for examples.

<!-- grade_id: validate_tail_number_answer, points: 0, grade: false, solution: true -->
```python
def validate_tail_number(n_number):
    '''
    Takes a string.
    Returns True if the string is a valid tail number.
    Returns False if the string is not a valid tail number.
    '''
    # YOUR CODE HERE
```

<!-- grade_id: validate_tail_number_test, points: 10, grade: true, solution: false -->
```python
assert_equal(validate_tail_number('N'), False) # N-number may only consist of one to five characters

assert_equal(validate_tail_number('N1'), True)
assert_equal(validate_tail_number('N99'), True)

assert_equal(validate_tail_number('NA'), False) # must start with a digit other than zero
assert_equal(validate_tail_number('N0'), False) # must start with a digit other than zero

assert_equal(validate_tail_number('N1A'), True)
assert_equal(validate_tail_number('N9Z'), True)

assert_equal(validate_tail_number('N1I'), False) # may not contain the letter I
assert_equal(validate_tail_number('N1O'), False) # may not contain the letter O

assert_equal(validate_tail_number('N1AA'), True)
assert_equal(validate_tail_number('N9ZZ'), True)
assert_equal(validate_tail_number('N1AAA'), False) # cannot end in a run of more than two letters
assert_equal(validate_tail_number('N1AI'), False) # may not contain the letter I
assert_equal(validate_tail_number('N1AO'), False) # may not contain the letter O

assert_equal(validate_tail_number('N10A'), True)
assert_equal(validate_tail_number('N99Z'), True)
assert_equal(validate_tail_number('N10I'), False) # may not contain the letter I
assert_equal(validate_tail_number('N11O'), False) # may not contain the letter O

assert_equal(validate_tail_number('N10AA'), True)
assert_equal(validate_tail_number('N99ZZ'), True)
assert_equal(validate_tail_number('N10AAA'), False) # cannot end in a run of more than two letters

assert_equal(validate_tail_number('N100'), True)
assert_equal(validate_tail_number('N999'), True)

assert_equal(validate_tail_number('N100A'), True)
assert_equal(validate_tail_number('N999Z'), True)
assert_equal(validate_tail_number('N100I'), False) # may not contain the letter I
assert_equal(validate_tail_number('N999O'), False) # may not contain the letter O

assert_equal(validate_tail_number('N100AA'), True)
assert_equal(validate_tail_number('N999ZZ'), True)
assert_equal(validate_tail_number('N100ZZZ'), False) # cannot end in a run of more than two letters

assert_equal(validate_tail_number('N1000'), True)
assert_equal(validate_tail_number('N9999'), True)

assert_equal(validate_tail_number('N1000A'), True)
assert_equal(validate_tail_number('N9999Z'), True)
assert_equal(validate_tail_number('N1000I'), False) # may not contain the letter I
assert_equal(validate_tail_number('N9999O'), False) # may not contain the letter O

assert_equal(validate_tail_number('N10000'), True)
assert_equal(validate_tail_number('N99999'), True)

assert_equal(validate_tail_number('N100000'), False) # N-number may only consist of one to five characters
assert_equal(validate_tail_number('N10000A'), False) # N-number may only consist of one to five characters
```