Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).

Make sure you fill in any place that says YOUR CODE HERE.
Do not write your answer in anywhere else other than where it says YOUR CODE HERE.

First, write your name and NetID below:

```python
NAME = 'WRITE YOUR NAME HERE'
NETID = 'WRITE YOUR NETID HERE'
```

# Problem 4.1. Simple Statistics.

In this problem, you will write a function that calculates the minimum, maximum, mean, and median values from a list of integers. Before writing the final function, I will break the problem into small steps so the autograder can give partial credit. In practice, you would probably write just one function `get_stats()` and be done.

## Function: get\_minimum()

- In the following code cell, write a function named `get_minimum()` that takes a list and returns the minimum value of that list. The return value should be an integer.

Tips:

- You might want to use [sorted()](https://docs.python.org/3/library/functions.html#sorted) to handle an unordered list.
- You may ignore any tips as long as your function does what it's supposed to, i.e. takes the specified argument(s) of the correct type(s) and returns the right values of the correct type(s). If your function passes the unit tests, your function is correct no matter which method you use.

<!-- grade_id: get_minimum_answer, points: 0, grade: false, solution: true -->
```python
def get_minimum(input_list):
    '''
    Accepts a list of integers, and returns the minimum of the list.
    
    Parameters
    ----------
    input_list: A list of integers.
      
    Returns
    -------
    minimum: An integer. The minimum of input_list.
    '''
    
    # YOUR CODE HERE
    
    return minimum
```

<!-- grade_id: get_minimum_test, points: 6, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(get_minimum([0, 1, 2, 3, 4]), 0)
assert_equal(get_minimum([1, 2, 3, 4]), 1)
assert_equal(get_minimum([4, 3, 2, 1, 0]), 0)
assert_equal(get_minimum([4, 1, 0, 2, 3]), 0)
assert_equal(get_minimum(list(range(100))), 0)
assert_equal(get_minimum(list(range(1, 100))), 1)
```

## Function: get\_maximum()

- In the following code cell, write a function named `get_maximum()` that takes a list and returns the maxium value of that list. The return value should be an integer.

<!-- grade_id: get_maximum_answer, points: 0, grade: false, solution: true -->
```python
def get_maximum(input_list):
    '''
    Accepts a list of integers, and returns the maximum of the list.
    
    Parameters
    ----------
    input_list: A list of integers.
      
    Returns
    -------
    maximum: An integer. The maximum of input_list.
    '''
    
    # YOUR CODE HERE
    
    return maximum
```

<!-- grade_id: get_maximum_test, points: 6, grade: true, solution: false -->
```python
assert_equal(get_maximum([0, 1, 2, 3, 4]), 4)
assert_equal(get_maximum([0, 1, 2, 3]), 3)
assert_equal(get_maximum([4, 3, 2, 1, 0]), 4)
assert_equal(get_maximum([3, 4, 0, 2, 1]), 4)
assert_equal(get_maximum(list(range(100))), 99)
```

## Function: get\_mean()

- In the following code cell, write a function named `get_mean()` that takes a list and returns the mean value of that list. The return value should be a float.

<!-- grade_id: get_mean_answer, points: 4, grade: false, solution: true -->
```python
def get_mean(input_list):
    '''
    Accepts a list of integers, and returns the mean of the list.
    
    Parameters
    ----------
    input_list: A list of integers.
      
    Returns
    -------
    mean: A float. The mean of input_list.
    '''
    
    # YOUR CODE HERE
    
    return mean
```

<!-- grade_id: get_mean_test, points: 6, grade: true, solution: false -->
```python
assert_equal(get_mean([0, 1, 2, 3]), 1.5)
assert_equal(get_mean([0, 1, 2, 3, 4]), 2.0)
assert_equal(get_mean([2, 0, 3, 1]), 1.5)
assert_equal(get_mean([1, 3, 4, 2, 0]), 2.0)
assert_equal(get_mean(list(range(0, 101))), 50.0)
assert_equal(get_mean(list(range(1, 100))), 50.0)
```

## Function: get\_median()

- In the following code cell, write a function named `get_median()` that takes a list and returns the median value of that list. The return value should be a float.

Tips:

- If there is an even number of values in the list, there is no single middle value; in this case, take the median to be the mean of the two middle values.

<!-- grade_id: get_median_answer, points: 4, grade: false, solution: true -->
```python
def get_median(input_list):
    '''
    Accepts a list of integers, and returns the median of the list.
    
    Parameters
    ----------
    input_list: A list of integers.
      
    Returns
    -------
    median: A float. The median of input_list.
            If there is an even number of values in the list,
            there is no single middle value.
            In this case, take the median to be the mean of the two middle values.

    '''
    
    # YOUR CODE HERE
    
    return median
```

<!-- grade_id: get_median_test, points: 6, grade: true, solution: false -->
```python
assert_equal(get_median([0, 1, 2, 3, 4]), 2.0)
assert_equal(get_median([4, 1, 3, 2, 0]), 2.0)
assert_equal(get_median([0, 1, 2, 3, 4, 5]), 2.5)
assert_equal(get_median([3, 4, 2, 1, 5, 0]), 2.5)
assert_equal(get_median(list(range(0, 101))), 50.0)
assert_equal(get_median(list(range(1, 101))), 50.5)
```

## Function: get\_stats()

- Write a function named `get_stats()` takes one argument, which should be a list of integers, and returns a tuple `(minimum, maximum, mean, median)`. The minimum and maximum values can be integers, but mean and median must be returned as floats.

<!-- grade_id: get_stats_answer, points: 0, grade: false, solution: true -->
```python
def get_stats(input_list):
    '''
    Accepts a list of integers, and returns a tuple of four numbers:
    minimum(int), maximum(int), mean(float), and median(float)
    
    Parameters
    ----------
    input_list: A list of integers.
      
    Returns
    -------
    minimum: An integer. The minimum of input_list.
    maximum: An integer. The maximum of input_list.
    mean: A float. The mean of input_list.
    median: A float. The median of input_list.
            If there is an even number of values in the list,
            there is no single middle value.
            In this case, take the median to be the mean of the two middle values.
    '''
    
    # YOUR CODE HERE

    return a_tuple
```

<!-- grade_id: get_stats_test, points: 6, grade: true, solution: false -->
```python
assert_equal(get_stats(list(range(51))), (0, 50, 25.0, 25.0))
assert_equal(get_stats(list(range(52))), (0, 51, 25.5, 25.5))

from random import shuffle
test_list = list(range(51))
shuffle(test_list)
assert_equal(get_stats(test_list), (0, 50, 25.0, 25.0))

test_list = list(range(52))
shuffle(test_list)
assert_equal(get_stats(test_list), (0, 51, 25.5, 25.5))
```