Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).
Make sure you fill in any place that says YOUR CODE HERE. Do not write your answer in anywhere else other than where it says YOUR CODE HERE.

First, write your name and NetID below:

```python
NAME = 'WRITE YOUR NAME HERE'
NETID = 'WRITE YOUR NETID HERE'
```

# Week 14 Assignment

In this assignment, we will visualize the total number of flights and current temperature of the top 20 airports using the 2001 flight data. This week's assignment is one long, continuous problem, but I split it up into two sections for easier grading. Before you start coding, read the entire notebook first to understand the big picture.

```python
import numpy as np
import pandas as pd
import sqlite3 as sl
import requests
from bs4 import BeautifulSoup

from nose.tools import assert_equal

from IPython.display import SVG, display_svg
```

This is where we are going:

```python
# run this cell if there's no image
r = requests.get("https://raw.githubusercontent.com/UI-DataScience/info490-fa15/master/Week14/assignment/top20.svg")
SVG(r.content)
```

The circles are the top 20 airports in the U.S. The size of each circle is proportional to the total number of arrivals and departures in 2001. The redder the color, the higher the temperature; the bluer the color, the lower the temperature. Thus, we will visualize three pieces of information in one single plot: the location of major airports, the size of the airports, and the current temperature.

## Problem 14.1. Top 20 Airports.

Recall that in [Problem 9.1](https://github.com/UI-DataScience/info490-fa15/blob/master/Week9/assignment/xml.ipynb) we wrote a function named `get_total_flights()` that adds the number of departures and the number of arrivals in 2001:

```python
dest_origin = pd.read_csv("2001.csv", encoding='latin-1', usecols=('Dest', 'Origin'))
flights = get_total_flights(dest_origin)
```

And we also found which 20 airports had the most number of flights:

```python
top20 = flights.sort(ascending=False, inplace=False)[:20]
```

Suppose that we have stored the result (the data frame `top20`) in an SQL database named [top20.db](https://github.com/UI-DataScience/info490-fa15/blob/master/Week14/assignment/top20.db?raw=true). You can download `top20.db` from the course repository on Github:

```python
!wget -O top20.db https://github.com/UI-DataScience/info490-fa15/blob/master/Week14/assignment/top20.db?raw=true
```

### Function: read\_top20()

Your first task is to

- Write a functoin named `read_top20()` that takes the file name (`top20.db`) and returns a Pandas data frame with the contents of the `top20` table in the database.

<!-- grade_id: read_top20_answer, points: 0, grade: false, solution: true -->
```python
def read_top20(db):
    """
    Takes the file name of an SQL database.
    Returns a Pandas data frame of all rows in the "top20" table of that database.
    
    Parameters
    ----------
    db (str): file name of SQL database.
    
    Returns
    -------
    A Pandas.DataFrame with all rows in "top20" table.
    """
    
    ## YOUR CODE HERE
    
    return result
```

So, when we do

```python
top20 = read_top20("top20.db")
print(top20)
```

the output should be

```
   iata  flights
0   ORD   682636
1   DFW   624361
2   ATL   503163
3   LAX   450019
4   PHX   368631
5   STL   324477
6   DTW   297522
7   MSP   284955
8   LAS   272293
9   BOS   266032
10  DEN   265184
11  IAH   257193
12  CLT   256626
13  SFO   243473
14  EWR   241016
15  PHL   239390
16  LGA   232964
17  PIT   212738
18  SEA   205486
19  BWI   199674
```

```python
top20 = read_top20("top20.db")
print(top20)
```

Make sure that your function passes the following unit test.

<!-- grade_id: read_top20_test, points: 10, grade: true, solution: false -->
```python
answer = {
    "iata": [
        'ORD', 'DFW', 'ATL', 'LAX', 'PHX',
        'STL', 'DTW', 'MSP', 'LAS', 'BOS',
        'DEN', 'IAH', 'CLT', 'SFO', 'EWR',
        'PHL', 'LGA', 'PIT', 'SEA', 'BWI'],
    "flights": [
        682636, 624361, 503163, 450019, 368631,
        324477, 297522, 284955, 272293, 266032,
        265184, 257193, 256626, 243473, 241016,
        239390, 232964, 212738, 205486, 199674]
}

answer_df = pd.DataFrame(answer)

np.testing.assert_array_equal(top20["iata"].values, answer_df["iata"].values)
np.testing.assert_array_equal(top20["flights"].values, answer_df["flights"].values)
```

Note that our airports are identified by IATA codes in the dataframe. However, the circles in [airports.svg](https://raw.githubusercontent.com/UI-DataScience/info490-fa15/master/Week14/assignment/airports.svg) are identified by city names, e.g., "Chicago" or "San Francisco". For example, the circle that corresponds to O'Hare Airport is represented by the following XML code in [airports.svg](https://raw.githubusercontent.com/UI-DataScience/info490-fa15/master/Week14/assignment/airports.svg):
  
```xml
<circle cx="367.68198" cy="118.095"
    stroke="black" stroke-width="1" r="5" fill="red"
    id="Chicago"/>
```

So we need to match the IATA codes with the city names.
  
Note that this problem is similar to `get_airport()` in [Problem 5.2](https://github.com/UI-DataScience/info490-fa15/blob/master/Week5/assignment/requests.ipynb) and `is_delayed()` in [Problem 9.2](https://github.com/UI-DataScience/info490-fa15/blob/master/Week9/assignment/xml.ipynb), but you should use [Beautiful Soup](http://www.crummy.com/software/BeautifulSoup/bs4/doc/) in this problem.


### Function: get\_city()

- Write a function named `get_city` that takes an XML (str) and returns the city name (str). See the unit tests below.

<!-- grade_id: get_city_answer, points: 0, grade: false, solution: true -->
```python
def get_city(xml):
    """
    Takes an XML and returns the city name.
    
    Parameters
    ----------
    xml (str): An XML script.
    
    Returns
    -------
    A string.
    """
    
    ## YOUR CODE HERE
    
    return result
```

The following code cell tests the function `get_city()`.

<!-- grade_id: get_city_test, points: 10, grade: true, solution: false -->
```python
test1 = '''
<?xml version="1.0" encoding="UTF-8"?><AirportStatus> 
					<Delay>true</Delay>
				
					<IATA>ORD</IATA>
				
					<State>Illinois</State>
				
					<Name>Chicago OHare International</Name>
				
					<Weather><Visibility>10.00</Visibility>
				
					<Weather>Partly Cloudy</Weather>
				
					<Meta><Credit>NOAA&apos;s National Weather Service</Credit>
				
					<Updated>12:51 PM Local</Updated>
				
					<Url>http://weather.gov/</Url></Meta>
				
					<Temp>53.0 F (11.7 C)</Temp>
				
					<Wind>Northwest at 12.7mph</Wind></Weather>
				
					<ICAO>KORD</ICAO>
				
					<City>Chicago</City>
				
					<Status><Reason>VOL:Multi-taxi</Reason>
				
					<ClosureBegin></ClosureBegin>
				
					<EndTime></EndTime>
				
					<MinDelay>16 minutes</MinDelay>
				
					<AvgDelay></AvgDelay>
				
					<MaxDelay>30 minutes</MaxDelay>
				
					<ClosureEnd></ClosureEnd>
				
					<Trend>Increasing</Trend>
				
					<Type>Departure</Type></Status>
				</AirportStatus> 
'''.strip()

test2 = '''
<?xml version="1.0" encoding="UTF-8"?><AirportStatus> 
					<Delay>false</Delay>
				
					<IATA>SFO</IATA>
				
					<State>California</State>
				
					<Name>San Francisco International</Name>
				
					<Weather><Visibility>10.00</Visibility>
				
					<Weather>Partly Cloudy</Weather>
				
					<Meta><Credit>NOAA&apos;s National Weather Service</Credit>
				
					<Updated>12:56 PM Local</Updated>
				
					<Url>http://weather.gov/</Url></Meta>
				
					<Temp>68.0 F (20.0 C)</Temp>
				
					<Wind>North at 5.8mph</Wind></Weather>
				
					<ICAO>KSFO</ICAO>
				
					<City>San Francisco</City>
				
					<Status><Reason>No known delays for this airport.</Reason>
				
					<ClosureBegin></ClosureBegin>
				
					<EndTime></EndTime>
				
					<MinDelay></MinDelay>
				
					<AvgDelay></AvgDelay>
				
					<MaxDelay></MaxDelay>
				
					<ClosureEnd></ClosureEnd>
				
					<Trend></Trend>
				
					<Type></Type></Status>
				</AirportStatus> 
'''.strip()

assert_equal(get_city(test1), "Chicago")
assert_equal(get_city(test2), "San Francisco")
```

### Function: get\_temp()

- Write a function named `get_temp` that takes an XML (str) and returns the current temperature (float).

<!-- grade_id: get_temp_answer, points: 0, grade: false, solution: true -->
```python
def get_temp(xml):
    """
    Takes an XML and returns the temperature.
    
    Parameters
    ----------
    xml (str): An XML script.
    
    Returns
    -------
    A float.
    """
    
    ## YOUR CODE HERE
    
    return result
```

The following code cell tests the function `get_temp()`.

<!-- grade_id: get_temp_test, points: 10, grade: true, solution: false -->
```python
assert_equal(get_temp(test1), 53.0)
assert_equal(get_temp(test2), 68.0)
```

Let's use `get_city()` and `get_temp()` to add the city names and temperatures to the data frame. Make sure that `get_city()` and `get_temp()` passed the unit tests before running the following code cell, because too many HTTP requests may lock you out for a while. You should get

```python
print(top20)
```
```
   iata  flights             city  temperature
0   ORD   682636          Chicago           53
1   DFW   624361  Dallas-Ft Worth           81
2   ATL   503163          Atlanta           75
3   LAX   450019      Los Angeles           71
4   PHX   368631          Phoenix           91
5   STL   324477         St Louis           63
6   DTW   297522          Detroit           53
7   MSP   284955      Minneapolis           44
8   LAS   272293        Las Vegas           76
9   BOS   266032           Boston           61
10  DEN   265184           Denver           61
11  IAH   257193          Houston           86
12  CLT   256626        Charlotte           74
13  SFO   243473    San Francisco           66
14  EWR   241016           Newark           63
15  PHL   239390     Philadelphia           65
16  LGA   232964         New York           61
17  PIT   212738       Pittsburgh           55
18  SEA   205486          Seattle           58
19  BWI   199674        Baltimore           63
```

```python
for idx, row in top20.iterrows():
    url = "https://raw.githubusercontent.com/UI-DataScience/info490-fa15/master/Week14/assignment/xml/{}"
    r = requests.get(url.format(row["iata"]))
    top20.loc[idx, "city"] = get_city(r.text)
    top20.loc[idx, "temperature"] = get_temp(r.text)
    
print(top20)
```

## Problem 14.2. Data Visualization.

- Use the blank U.S. map with 20 major cities
  [airports.svg](https://rawgit.com/INFO490/spring2015/master/week09/images/airports.svg)
  to visualize the `top20` dataframe.
  
Run the following cell to display the SVG template..

```python
resp = requests.get("https://raw.githubusercontent.com/UI-DataScience/info490-fa15/master/Week14/assignment/airports.svg")
usairports = resp.content
display_svg(usairports, raw=True)
```

You should use BeautifulSoup to adjust the size and color of the circles to create a figure similar to the following. Again, the radius of a circle is proportional to the total traffic of that airport. The color of a circle depends on the current temperature at that airport.

```python
resp = requests.get("https://raw.githubusercontent.com/UI-DataScience/info490-fa15/master/Week14/assignment/top20.svg")
SVG(resp.content)
```

(Your figure does not have to look exactly the same. Choose a different set of sizes and colors if you wish.)

[Lesson 2](https://github.com/UI-DataScience/info490-fa15/blob/master/Week14/lesson2.md) and [Lesson 3](https://github.com/UI-DataScience/info490-fa15/blob/master/Week14/lesson3.md) cover how to use BeautifulSoup to parse an SVG file (which is in XML format), search for the FIPS code of each county, and replace the colors using a dataframe (unemployment data). Although we are searching for different objects (circles) and using a different data set (`top20`), the same principle applies.

Remember that each circle is represented by the following XML code:

```
<circle cx="367.68198" cy="118.095"
    stroke="black" stroke-width="1" r="5" fill="red"
    id="Chicago"/>
```

So you should use BeautifulSoup to parse the `usairports` (the contents of `airports.svg` file) and search for all instances of `circle`. Then, iterate over each instance of circle and replace the `r` variable with the number of flights flights in the `top20` dataframe (multiplied/divided by some proportionaly constant since a circle with a radius `r="682636"` won't work). You should also replace the `fill` variable with a color the temperature `temp`. After replacing all instances of `circle`, display the result with `display_svg(soup.prettify(), raw=True)`.

Write your answer for Problem 14.2 in the following code cell.

<!-- grade_id: svg, points: 30, grade: true, solution: true -->
```python
## YOUR CODE HERE
```