Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).

Make sure you fill in any place that says YOUR CODE HERE.
Do not write your answer in anywhere else other than where it says YOUR CODE HERE.

First, write your name and NetID below:

```python
NAME = 'WRITE YOUR NAME HERE'
NETID = 'WRITE YOUR NETID HERE'
```

# Problem 9.1. IPython notebook as JSON.

```python
import os
import json
import requests

from nose.tools import assert_equal
```

IPython/Jupyter notebooks are actually simple JSON documents. You may have noticed this if you are a Mac user, as OS X tries to attach `.json` at the end of the file name when you try to download an `.ipynb` file to your hard disk. It's because all information in the notebook is stored as JSON texts.

See [The Jupyter Notebook Format](https://ipython.org/ipython-doc/3/notebook/nbformat.html) for details.

In this problem, we will use the `json` module to parse a notebook file and extract some information about the notebook. We will use [intro2ipy.ipynb](https://github.com/UI-DataScience/info490-fa15/blob/master/Week2/notebooks/intro2ipy.ipynb) from Week 2 as an example.

```python
r = requests.get('https://raw.githubusercontent.com/UI-DataScience/info490-fa15/master/Week2/notebooks/intro2ipy.ipynb')

with open('intro2ipy.ipynb', 'w') as f:
    f.write(r.text)
```

## Function: get\_keys()

- Write a function named `get_keys()` that takes a file name (str) and returns the keys (list of strings) of the dictionary at the [top level](https://ipython.org/ipython-doc/3/notebook/nbformat.html#top-level-structure).

In other words, the function opens the JSON file, reads the file as a Python dictionary, and returns the keys of that dictionary. The dictionary you get from reading the JSON file is nested (there's a dictionary inside a dictionary inside a dictionary), and you should return the keys of the outermost (the top-level) dictionary.

<!-- grade_id: get_keys_answer, points: 0, grade: false, solution: true -->
```python
def get_keys(filename):
    '''
    Takes the file name (str) of a JSON text file and returns the keys of the top-level dictionary.
    
    Parameters
    ----------
    filename (str): a JSON file.
    
    Returns
    -------
    A list of strings.
    '''
    
    # YOUR CODE HERE
    
    return result
```

As shown in the [documentation](https://ipython.org/ipython-doc/3/notebook/nbformat.html#top-level-structure), you should get

```
['cells', 'nbformat_minor', 'metadata', 'nbformat']
```

when you call `get_keys()` for a typical notebook file. Run the following code cell to confirm this visually.

```python
print(get_keys('intro2ipy.ipynb'))
```

Your score is determined by the following tests.

<!-- grade_id: get_keys_test, points: 6, grade: true, solution: false -->
```python
test1 = get_keys('intro2ipy.ipynb')
answer1 = ['cells', 'nbformat_minor', 'metadata', 'nbformat']

assert_equal(len(test1), len(answer1))
assert_equal(set(test1), set(answer1))

test2 = {
    'A': 1,
    'B': {'C': 2, 'D': 3},
    'C': {
        'E': {'F': 4},
        'G': {'H': 5, 'I': 6}
    }
}

answer2 = ['A', 'B', 'C']

with open('test.json', 'w') as f:
    json.dump(test2, f)

assert_equal(len(get_keys('test.json')), len(answer2))
assert_equal(set(get_keys('test.json')), set(answer2))

os.remove('test.json')
```

## Function: get\_version()

Verion information is always important. For example, the course docker image runs Python 3.4, while the latest version of Python 3 is version 3.5, so new features in Python 3.5 may or may not work in 3.4. Or, if you try to run a Python 3 notebook using a Python 2 kernel, it will throw errors. Furthermore, IPython notebooks themselves have different format versions. We are using "nbformat" 4, and if you try to run this notebook in older versions of IPython kernels that only recognizes nbformats 3 or less, it won't run.

- Write a function named `get_version()` that takes a file name (str) and returns a tuple of ("the programming language of the kernel", "the version of the language", nbformat). The data type of this tuple is `(str, str, int)`.

<!-- grade_id: get_version_answer, points: 0, grade: false, solution: true -->
```python
def get_version(filename):
    '''
    Takes a file name (str) of a JSON file.
    Returns a tuple of ("the programming language of the kernel", "the version of the language", nbformat).
    
    Parameters
    ----------
    filename (str): a JSON file.
    
    Returns
    -------
    A tuple of (str, str, int)
    '''
    
    # YOUR CODE HERE
    
    return name, version, nbformat
```

```python
print(get_version('intro2ipy.ipynb'))
```

Make sure that the following code cell does not produce an error.

<!-- grade_id: get_version_test, points: 8, grade: true, solution: false -->
```python
test1 = get_version('intro2ipy.ipynb')
answer1 = ('python', '3.4.0', 4)

assert_equal(test1, answer1)

test2 = {
  "metadata" : {
    "signature": "hex-digest", # used for authenticating unsafe outputs on load
    "kernel_info": {
        # if kernel_info is defined, its name field is required.
        "name" : "the name of the kernel"
    },
    "language_info": {
        # if language_info is defined, its name field is required.
        "name" : "the programming language of the kernel",
        "version": "the version of the language",
        "codemirror_mode": "The name of the codemirror mode to use [optional]"
    }
  },
  "nbformat": 4,
  "nbformat_minor": 0,
  "cells" : [
      # list of cell dictionaries, see below
  ],
}

answer2 = ("the programming language of the kernel", "the version of the language", 4)

with open('test.json', 'w') as f:
    json.dump(test2, f)
    
assert_equal(get_version('test.json'), answer2)

os.remove('test.json')
```

## Function: count\_code\_cells()

- Write a function named `count_code_cells()` that takes a filename (str) and a cell type (str), and returns the number count of that cell type (int).

<!-- grade_id: count_code_cells_answer, points: 0, grade: false, solution: true -->
```python
def count_code_cells(filename, cell_type):
    '''
    Takes a filename and a cell type, and returns the number count of that cell type.
    
    Parameters
    ----------
    filename (str): a JSON file.
    cell_type (str): "code", "markdown", etc.
    
    Returns
    -------
    An int.
    '''

    # YOUR CODE HERE

    return count
```

```python
n_code = count_code_cells('intro2ipy.ipynb', 'code')
n_markdown = count_code_cells('intro2ipy.ipynb', 'markdown')

print('There are {} code cells and {} markdown cells.'.format(n_code, n_markdown))
```

<!-- grade_id: count_code_cells_test, points: 8, grade: true, solution: false -->
```python
assert_equal(count_code_cells('intro2ipy.ipynb', 'code'), 5)
assert_equal(count_code_cells('intro2ipy.ipynb', 'markdown'), 16)

test = {
  "cells" : [
    {
      "cell_type" : "type1",
      "metadata" : {},
      "source" : "single string or [list, of, strings]",
      },
    {
      "cell_type" : "type1",
      "metadata" : {},
      "source" : "single string or [list, of, strings]",
      },
    {
      "cell_type" : "type2",
      "metadata" : {},
      "source" : "single string or [list, of, strings]",
      }
  ],
} 

with open('test.json', 'w') as f:
    json.dump(test, f)

assert_equal(count_code_cells('test.json', 'type1'), 2)
assert_equal(count_code_cells('test.json', 'type2'), 1)
assert_equal(count_code_cells('test.json', 'type3'), 0)

os.remove('test.json')
```

## Function: get\_nth\_code\_cell()

- Write a function named `get_nth_code_cell()` that takes a file name (str) and an integer `n`, and returns the **source code** in the `n`-th code cell (a list of strings for each line in the source code). Note that the `cell_type` should be "code" (not "markdown").

See below for examples and compare with [intro2ipy.ipynb](https://github.com/UI-DataScience/info490-fa15/blob/master/Week2/notebooks/intro2ipy.ipynb).

<!-- grade_id: get_nth_code_cell_answer, points: 0, grade: false, solution: true -->
```python
def get_nth_code_cell(filename, n):
    '''
    Takes a file name (str) and an integer n.
    Returns source code of that n-th code cell.
    
    Parameters
    ----------
    filename (str): a JSON file.
    n (int): 
    
    Returns
    -------
    A list of strings.
    '''

    # YOUR CODE HERE
    
    return result
```

```python
print(get_nth_code_cell('intro2ipy.ipynb', 0)) # change n
```

<!-- grade_id: get_nth_code_cell_test, points: 8, grade: true, solution: false -->
```python
answer0 = ['%lsmagic']
answer1 = ['print("Hello World!")']
answer2 = [
    'import numpy as np\n',
   'from scipy.integrate import quad\n',
   '\n',
   'print("The Integral = %3.1f" % quad(np.sin, -2.0 * np.pi, np.pi)[0])'
]

assert_equal(get_nth_code_cell('intro2ipy.ipynb', 0), answer0)
assert_equal(get_nth_code_cell('intro2ipy.ipynb', 1), answer1)
assert_equal(get_nth_code_cell('intro2ipy.ipynb', 2), answer2)

test = {
  "cells" : [
    {
      "cell_type" : "code",
      "metadata" : {},
      "source" : ["first code"],
      },
    {
      "cell_type" : "markdown",
      "metadata" : {},
      "source" : ["first markdown"],
      },
    {
      "cell_type" : "markdown",
      "metadata" : {},
      "source" : ["second markdown"],
      },
    {
      "cell_type" : "code",
      "metadata" : {},
      "source" : ["second code", "third code"],
      }
  ],
} 

with open('test.json', 'w') as f:
    json.dump(test, f)
    
assert_equal(get_nth_code_cell('test.json', 0), ['first code'])
assert_equal(get_nth_code_cell('test.json', 1), ['second code', 'third code'])

os.remove('test.json')
```