Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).

Make sure you fill in any place that says YOUR CODE HERE.
Do not write your answer in anywhere else other than where it says YOUR CODE HERE.

First, write your name and NetID below:

```python
NAME = 'WRITE YOUR NAME HERE'
NETID = 'WRITE YOUR NETID HERE'
```

# Problem 3.1. A famous simple interview question.

This problem is a famous programming job interview question designed to filter out those who think they can program but in reality can't. If you can easily solve this problem, you are already better than 99% of job candidates who apply for a programming job. The problem statement is simple:

- Print every number from 1 to 100.
- If the number is a multiple of 3, print "info" instead.
- If the number is a multiple of 5, print "matics" instead.
- If the number is a multiple of both 3 and 5, print "informatics" instead.

Note the spelling of "informatics". There's an R in informatics.

## 1. Function: get_informatics()

Use conditionals and the `str()` function to write a function named `get_informatics()` that takes an integer and returns a string. I emphasize that the return value is of type **str**, not int.

To convert an integer type into a string, you can use [str()](https://docs.python.org/3/library/stdtypes.html#str). For example,

```python
x = 1
print("x is a type {}.".format(type(x)))
```
```shell
<class 'int'>
```
```python
y = str(x)
print("y is a type {}.".format(type(y)))
```
```shell
<class 'str'>
```

<!-- grade_id: get_informatics, points: 0, grade: false, solution: true -->
```python
def get_informatics(number):
    '''
    Prints every number from 1 to 100. But for multiples of three, print "info" instead;
    and for multiples of five, print "matics." For numbers which are mltiples of both
    three and five, print "informatics."
    
    Parameters
    __________
    input: An integer greater than 0.
    
    Returns
    _______
    output: A string.
    '''
    
    # YOUR CODE HERE
    
    return result
```

If the following code cell doesn't produce any error, it means that your code works.

<!-- grade_id: test_get_informatics, points: 20, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(get_informatics(1), '1')
assert_equal(get_informatics(2), '2')
assert_equal(get_informatics(3), 'info')
assert_equal(get_informatics(4), '4')
assert_equal(get_informatics(5), 'matics')
assert_equal(get_informatics(6), 'info')
assert_equal(get_informatics(7), '7')
assert_equal(get_informatics(8), '8')
assert_equal(get_informatics(9), 'info')
assert_equal(get_informatics(10), 'matics')
assert_equal(get_informatics(11), '11')
assert_equal(get_informatics(12), 'info')
assert_equal(get_informatics(13), '13')
assert_equal(get_informatics(14), '14')
assert_equal(get_informatics(15), 'informatics')
assert_equal(get_informatics(16), '16')
assert_equal(get_informatics(17), '17')
assert_equal(get_informatics(18), 'info')
assert_equal(get_informatics(19), '19')
assert_equal(get_informatics(20), 'matics')
assert_equal(get_informatics(21), 'info')
assert_equal(get_informatics(22), '22')
assert_equal(get_informatics(23), '23')
assert_equal(get_informatics(24), 'info')
assert_equal(get_informatics(25), 'matics')
assert_equal(get_informatics(26), '26')
assert_equal(get_informatics(27), 'info')
assert_equal(get_informatics(28), '28')
assert_equal(get_informatics(29), '29')
assert_equal(get_informatics(30), 'informatics')
assert_equal(get_informatics(31), '31')
assert_equal(get_informatics(32), '32')
assert_equal(get_informatics(33), 'info')
assert_equal(get_informatics(34), '34')
assert_equal(get_informatics(35), 'matics')
assert_equal(get_informatics(36), 'info')
assert_equal(get_informatics(37), '37')
assert_equal(get_informatics(38), '38')
assert_equal(get_informatics(39), 'info')
assert_equal(get_informatics(40), 'matics')
assert_equal(get_informatics(41), '41')
assert_equal(get_informatics(42), 'info')
assert_equal(get_informatics(43), '43')
assert_equal(get_informatics(44), '44')
assert_equal(get_informatics(45), 'informatics')
assert_equal(get_informatics(46), '46')
assert_equal(get_informatics(47), '47')
assert_equal(get_informatics(48), 'info')
assert_equal(get_informatics(49), '49')
assert_equal(get_informatics(50), 'matics')
```

## 2. Iteration

In the following code cell, use the `get_informatics()` function and iteration (a `for` loop or a `while` loop) to print every number from 1 to 100. Multiples of 3 and 5 should be replaced according to the rules in `get_informatics()`. A sample output is provided after the code cell.

<!-- grade_id: print_100, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

```
1
2
info
4
matics
info
7
8
info
matics
11
info
13
14
informatics
16
17
info
19
matics
info
22
23
info
matics
26
info
28
29
informatics
31
32
info
34
matics
info
37
38
info
matics
41
info
43
44
informatics
46
47
info
49
matics
info
52
53
info
matics
56
info
58
59
informatics
61
62
info
64
matics
info
67
68
info
matics
71
info
73
74
informatics
76
77
info
79
matics
info
82
83
info
matics
86
info
88
89
informatics
91
92
info
94
matics
info
97
98
info
matics
```